[1501173586] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501173593] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501173623] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501173668] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501173684] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501173713] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501173714] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501173760] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501173803] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501173804] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501173834] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501173880] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501173923] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501173954] HOST ALERT: ip-10-0-0-123.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501174032] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501174043] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501174073] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501174163] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501174163] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501174193] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501174283] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501174313] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501174403] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501174433] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501174523] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501174553] HOST ALERT: ip-10-0-0-235.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179560] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179561] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179568] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179568] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179572] SERVICE ALERT: ip-10-0-7-103.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179579] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179579] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179582] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179584] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179589] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179590] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179591] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179595] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179596] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179598] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179598] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1501179602] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179612] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179612] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179614] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179616] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501179619] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179621] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179622] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179624] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179625] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179626] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179629] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179633] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179634] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179643] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179644] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179645] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179646] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179647] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179648] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179649] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179650] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179651] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179651] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179652] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179653] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179653] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179655] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179656] SERVICE ALERT: ip-10-0-7-103.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501179673] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179674] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179675] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179676] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179678] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179679] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179680] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179680] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179681] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179681] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501179682] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179685] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179686] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179688] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179688] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179692] SERVICE ALERT: ip-10-0-7-103.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179694] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179699] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179702] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179704] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179709] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179710] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179711] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179715] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179716] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179718] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179718] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179720] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179722] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179724] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179726] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179727] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179729] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501179730] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179731] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179732] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179736] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501179742] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179744] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179753] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179754] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179762] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179762] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179763] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179764] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179765] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179766] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179767] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179768] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179768] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179769] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179769] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179770] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501179770] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179771] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179771] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179772] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179773] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179773] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179775] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179776] SERVICE ALERT: ip-10-0-7-103.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501179793] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179794] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179795] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179796] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179798] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179799] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179800] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179801] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501179802] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179805] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179806] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179814] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179826] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179830] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179831] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179838] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179838] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179842] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179844] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179846] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179847] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179849] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501179850] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179851] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179861] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179862] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179864] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179873] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179874] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179882] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179887] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179888] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179889] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179890] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501179891] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179893] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179913] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179914] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179915] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179916] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179918] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179919] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179920] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179921] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501179922] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179923] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179925] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179926] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179946] HOST ALERT: ip-10-0-0-65.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179950] HOST ALERT: ip-10-0-1-165.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179951] HOST ALERT: ip-10-0-1-102.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179958] HOST ALERT: ip-10-0-2-167.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179958] HOST ALERT: ip-10-0-5-128.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179962] HOST ALERT: ip-10-0-7-103.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179964] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179966] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179967] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179969] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501179970] HOST ALERT: ip-10-0-3-44.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501179971] HOST ALERT: ip-10-0-3-94.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501180002] HOST ALERT: ip-10-0-1-47.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501180008] HOST ALERT: ip-10-0-2-193.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501180009] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501180010] HOST ALERT: ip-10-0-2-5.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%