[1503426002] SERVICE ALERT: ip-10-0-1-16.us-west-2.compute.internal;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 147.66 ms
[1503426122] SERVICE ALERT: ip-10-0-1-16.us-west-2.compute.internal;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1504041978] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041981] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041982] SERVICE ALERT: ip-10-0-0-173.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041983] SERVICE ALERT: ip-10-0-0-254.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041983] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041984] SERVICE ALERT: ip-10-0-0-63.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041985] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041985] SERVICE ALERT: ip-10-0-0-85.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041986] SERVICE ALERT: ip-10-0-1-11.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041986] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041987] SERVICE ALERT: ip-10-0-1-136.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041988] SERVICE ALERT: ip-10-0-1-16.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041988] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041989] SERVICE ALERT: ip-10-0-1-69.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041989] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041990] SERVICE ALERT: ip-10-0-1-89.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041991] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041991] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041992] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041993] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041993] SERVICE ALERT: ip-10-0-2-86.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041994] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041994] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041995] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041996] SERVICE ALERT: ip-10-0-3-56.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041996] SERVICE ALERT: ip-10-0-3-54.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504041997] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;Local Disk;CRITICAL;SOFT;1;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042098] SERVICE ALERT: ip-10-0-5-128.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042101] SERVICE ALERT: ip-10-0-0-123.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042102] SERVICE ALERT: ip-10-0-0-173.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042103] SERVICE ALERT: ip-10-0-0-254.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042103] SERVICE ALERT: ip-10-0-0-235.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042104] SERVICE ALERT: ip-10-0-0-63.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042105] SERVICE ALERT: ip-10-0-0-65.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042105] SERVICE ALERT: ip-10-0-0-85.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042106] SERVICE ALERT: ip-10-0-1-102.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042106] SERVICE ALERT: ip-10-0-1-11.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042107] SERVICE ALERT: ip-10-0-1-136.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042108] SERVICE ALERT: ip-10-0-1-16.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042108] SERVICE ALERT: ip-10-0-1-165.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042109] SERVICE ALERT: ip-10-0-1-47.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042109] SERVICE ALERT: ip-10-0-1-69.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042110] SERVICE ALERT: ip-10-0-1-89.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042111] SERVICE ALERT: ip-10-0-2-193.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042111] SERVICE ALERT: ip-10-0-2-167.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042112] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042113] SERVICE ALERT: ip-10-0-2-86.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042113] SERVICE ALERT: ip-10-0-2-5.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042114] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042114] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042115] SERVICE ALERT: ip-10-0-3-44.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042116] SERVICE ALERT: ip-10-0-3-56.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042116] SERVICE ALERT: ip-10-0-3-54.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory
[1504042117] SERVICE ALERT: ip-10-0-3-94.us-west-2.compute.internal;Local Disk;CRITICAL;HARD;2;DISK CRITICAL - 20% is not accessible: No such file or directory