[1501044288] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501044347] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501044407] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501044467] SERVICE ALERT: localhost;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 0 MB) - Swap is either disabled, not present, or of zero size.
[1501044848] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501044860] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501044875] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501044905] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501044966] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501044968] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501044995] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501044998] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501045046] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501045086] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501045118] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501045166] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501045206] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501045238] HOST ALERT: ip-10-0-7-149.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501085885] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;1;(Host check timed out after 31.01 seconds)
[1501085897] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501085920] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501085923] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501085953] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501085984] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501086017] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501086040] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501086043] SERVICE ALERT: ip-10-0-3-196.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501086073] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501086104] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501086160] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501086193] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501086224] HOST ALERT: ip-10-0-3-196.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501086684] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1501086714] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501086719] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1501086749] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1501086793] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1501086803] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1501086833] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1501086838] SERVICE ALERT: ip-10-0-3-26.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1501086868] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1501086913] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1501086953] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1501086988] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1501087033] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1501087073] HOST ALERT: ip-10-0-3-26.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1501103631] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Check SSH;UNKNOWN;HARD;2;Usage:
[1501105092] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout