[1504717612] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Total Mesos CPUS;CRITICAL;SOFT;1;CRITICAL: Mesos metrics: allocator/mesos/resources/cpus/total=84 (84>36, w=30/c=36)
[1504717672] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Total Mesos CPUS;CRITICAL;HARD;2;CRITICAL: Mesos metrics: allocator/mesos/resources/cpus/total=84 (84>36, w=30/c=36)
[1504718236] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;CPUS Used;CRITICAL;SOFT;1;CRITICAL: Mesos metrics: master/cpus_used=41.60 (41.60>40, w=30/c=40)
[1504718273] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Total Mesos CPUS;OK;HARD;2;OK: Mesos metrics: allocator/mesos/resources/cpus/total=84
[1504718295] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;CPUS Used;CRITICAL;HARD;2;CRITICAL: Mesos metrics: master/cpus_used=41.60 (41.60>40, w=30/c=40)
[1504720094] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;CPUS Used;WARNING;HARD;2;WARNING: Mesos metrics: master/cpus_used=30.60 (30.60>30, w=30/c=40)
[1504720215] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;CPUS Used;CRITICAL;HARD;2;CRITICAL: Mesos metrics: master/cpus_used=41.60 (41.60>40, w=30/c=40)
[1504721895] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;CPUS Used;OK;HARD;2;OK: Mesos metrics: master/cpus_used=42.60
[1504722514] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Task Lost;UNKNOWN;SOFT;1;UNKNOWN: metric not found 'master/task_lost'
[1504722573] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Task Lost;UNKNOWN;HARD;2;UNKNOWN: metric not found 'master/task_lost'
[1504728261] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Slave Inactive;WARNING;SOFT;1;WARNING: Mesos metrics: master/slaves_inactive=1 (1>0, w=0/c=2)
[1504728261] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Check SSH;CRITICAL;SOFT;1;Server answer:
[1504728320] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Slave Inactive;WARNING;HARD;2;WARNING: Mesos metrics: master/slaves_inactive=1 (1>0, w=0/c=2)
[1504728390] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;Check SSH;CRITICAL;HARD;2;CRITICAL - Socket timeout
[1504728395] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1504728420] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1504728480] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1504728515] SERVICE ALERT: ip-10-0-2-46.us-west-2.compute.internal;PING;CRITICAL;HARD;2;PING CRITICAL - Packet loss = 100%
[1504728530] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1504728600] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1504728620] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Slave Inactive;OK;HARD;2;OK: Mesos metrics: master/slaves_inactive=0
[1504728650] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1504728720] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1504728770] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1504728840] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1504728890] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1504728960] HOST ALERT: ip-10-0-2-46.us-west-2.compute.internal;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1504731690] SERVICE ALERT: ip-10-0-7-149.us-west-2.compute.internal;Task Lost;OK;HARD;2;OK: Mesos metrics: master/tasks_lost=24
[1504739160] SERVICE ALERT: ip-10-0-1-89.us-west-2.compute.internal;PING;WARNING;SOFT;1;PING WARNING - Packet loss = 0%, RTA = 100.43 ms
[1504739280] SERVICE ALERT: ip-10-0-1-89.us-west-2.compute.internal;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1504740909] SERVICE ALERT: ip-10-0-1-136.us-west-2.compute.internal;Free Mem;CRITICAL;SOFT;1;CRITICAL: Mesos metrics: system/mem_free_bytes=1948282880 (1948282880>50000000, w=100000000/c=50000000)
[1504740969] SERVICE ALERT: ip-10-0-1-136.us-west-2.compute.internal;Free Mem;CRITICAL;HARD;2;CRITICAL: Mesos metrics: system/mem_free_bytes=1949376512 (1949376512>50000000, w=100000000/c=50000000)